# Otter-Grader Demo

```python
def is_in_colab():
    try:
        import google.colab
        in_colab = True
    except ImportError:
        in_colab = False
    return in_colab

if is_in_colab():
    !git clone https://github.com/charlesfrye/otter-demo-labs.git
    %cd ./otter-demo-labs/
    !pip install otter-grader  # or requirements.txt
```

```python
import otter
grader = otter.Notebook(colab=is_in_colab())
```

**Question 1:** Write a function `square` that returns the square of its argument.

_Points:_ 2

```python
def square(x):
    ...
```

```python
grader.check("q1")
```

**Question 2:** Update the `foo.py` module to have a global variable named `BAR` that has value `"baz"`.

_Points:_ 1

```python
import foo
```

```python
grader.check("q2")
```

**Question 3:** Update the `foo.py` script to print `"Hello, world!"`.

_Points:_ 1

```python
...
# Once you've made changes in foo.py, this cell should print `Hello, world!`
!python foo.py
```

```python
grader.check("q3")
```

Tests for q1 (points: None)
```
>>> 
>>> square(1) == 1
True
```
failure_message: 1 ** 2 == 1 because 1 * 1 = 1
success_message: Good job!
```
>>> square(0) == 0
True
```
```
>>> square(2.5) == 6.25
True
```


Tests for q2 (points: None)
```
>>> 
>>> import foo
>>> 
>>> assert("bar") not in dir(foo)
```
failure_message: capitalization matters!
```
>>> 
>>> import foo
>>> 
>>> assert("BAR") in dir(foo)
```
failure_message: make sure you define BAR in foo.py
```
>>> 
>>> import foo
>>> 
>>> assert foo.BAR == "baz"
```
failure_message: check the value of foo.BAR
success_message: Good job!


Tests for q3 (points: None)
```
>>> 
>>> import subprocess
>>> 
>>> result = subprocess.run(["python", "foo.py"], capture_output=True)
>>> result.stdout
b'Hello, world!\n'
```
success_message: Good job!
