```python
# Initialize Otter
import otter
grader = otter.Notebook(colab=True)
```

# Otter-Grader Demo

```python
import pandas as pd
import numpy as np
%matplotlib inline
import otter
grader = otter.Notebook()
```

**Question 1:** Write a function `square` that returns the square of its argument.

_Points:_ 2

```python
def square(x):
    ...
```

```python
grader.check("q1")
```

**Question 2:** Update the `foo.py` module to have a global variable named `BAR` that has value `"baz"`.

_Points:_ 1

```python
del foo.BAR
```

```python
...

foo.BAR = ...
```

```python
grader.check("q2")
```

**Question 3:** Update the `foo.py` script to print `"Hello, world!"`.

_Points:_ 1

```python
!python foo.py
```

```python
grader.check("q2")
```

Tests for q1 (points: None)
```
>>> 
>>> square(1) == 1
True
```
failure_message: 1 ** 2 == 1 because 1 * 1 = 1
success_message: Good job!
```
>>> square(0) == 0
True
```


Tests for q2 (points: None)
```
>>> 
>>> import subprocess
>>> 
>>> result = subprocess.run(["python", "foo.py"], capture_output=True)
>>> result.stdout
b'Hello, world!\n'
```
success_message: Good job!
